---
title: Coding
description: 'easy examples u should solve without using google or anything else.'
---

## Vektoren erzeugen

```yaml
type: NormalExercise
key: a8ba0480e7
xp: 100
```



`@instructions`
Erzeugen Sie einen Vektor x, der aus den natürlichen Zahlen von 1 bis 10 besteht

`@hint`
c(x:y)

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
x <- c(1:10)
```

`@sct`
```{r}

```
